---
# Welcome to CE880
### This is your week-2 : Problem 1 notebook
### Spring Term 2022-23
#### Prepared by Dr Haider Raza

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function to pass a list of numbers and calculate the exponential of all elements in the list format. 
`Hint`: take input as list and inside the function covert it into numpy array then perform the function. `Example` : input `[1, 4, 8]` and and output should be:  `[0.0, 2.0, 3.0]` 

<!-- grade_id: myLog_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myExp(mylist):
    """Write a function to pass a list of numbers and calculate the exponential of all elements """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
in_array = [1, 2, 4, 6]
myExp(in_array)
```

<!-- grade_id: myLog_test, points: 1, grade: true, solution: false -->
```python
"""run these test cases"""
# Test case 1
my = [1, 2, 4, 6]
assert myExp(my) == [2.718281828459045, 7.38905609893065, 54.598150033144236, 403.4287934927351]

# Test case 2
my = [1, 10, 100]
assert myExp(my) == [2.718281828459045, 22026.465794806718, 2.6881171418161356e+43]
```

---

## Question 2:

Write a function to find the k number of the smallest values from a NumPy array and remove duplicate when returning as a list. 


For example: 
`nSmallest([23, 12, 1, 3, 1, 4, 0, 4, 5, 6],3) = [0, 1, 3]`

<!-- grade_id: myfibonacci_ans, points: 0, grade: false, solution: true -->
```python
def nSmallest(array, k):
    """Find the k number of the smallest values and remove duplicates"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
nSmallest([100, 200, 2000, 30, 110], 4)
```

<!-- grade_id: myfibonacci_test, points: 1, grade: true, solution: false -->
```python
"""run these test cases"""
assert nSmallest([23, 12, 1, 3, 1, 4, 0, 4, 5, 6], 3) == [0, 1, 3]
assert nSmallest([100, 200, 2000, 30, 110], 4) == [30, 100, 110, 200]
assert nSmallest([10, 10, 2, 5, 4, 5, 6, 8, 2], 2) == [2, 4]

```

---
## Question 3: 

 Write a program, which accept a list with `nan`s in it and count the maximum number of consecutive repeated `nan`.
 Example: In the `a = [nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.9999, 0.0001, 0.16, 0.101, nan, 0.16]` and answer is 2.

<!-- grade_id: max_nan_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
a = [nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.9999, 0.0001, 0.16, 0.101, nan, 0.16]
max_rep_nan(a)
```

<!-- grade_id: max_rep_test, points: 0, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""
a = [nan, nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 3

a = [nan, nan, nan, nan, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 5

a = [nan, 10, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 1

a = []
assert max_rep_nan(a) == 0
```

## Pandas Data Frame 


```python
from sklearn import datasets
import pandas as pd
```

```python
# import some data to play with
diabetes = datasets.load_diabetes()
X = pd.DataFrame(diabetes.data[:, :])  # we are taking all features.
y = diabetes.target
print(X)
print(y)
```

---
## Question 4:

Write a function to find the max number in the each column in the diabetes dataset and return the output as a list. Please make sure to round the output to two decimal points. 

<!-- grade_id: filter_1_ans, points: 0, grade: false, solution: true -->
```python
def my_max_min(X):
    """Write a function to find the max number in the each column and return 
    as a list."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: filter_1_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
assert my_max_min(X) == [0.11, 0.05, 0.17, 0.13, 0.15, 0.2, 0.18, 0.19, 0.13, 0.14]
```

---
## Question 5:


Write a function to calculate the `mean` and `standard deviation` of diabetes dataset (first column only). `Hint`: You can use NumPy aggregate functions. Create a list and append value to it and return the output. 

<!-- grade_id: stats_ans, points: 0, grade: false, solution: true -->
```python
def stats(X):
    """Write a functon to calculate mean and standard deviation of IRIS dataset (last column only)."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
stats(X)
```

<!-- grade_id: stats_test, points: 2, grade: true, solution: false -->
```python
"""Check your output"""
# Test case
x = stats(X)
y= [-3.6396225400041895e-16, 0.04761904761904939]
np.testing.assert_allclose(x, y, rtol=0.05, atol=0)
```